```python
# Initialize Otter
import otter
grader = otter.Notebook("hw-demo-nba-wnba.ipynb")
```

```python
from datascience import *
```

## HW Demo: WNBA and NBA Player Data: Exploring Datasets

<div class="alert alert-block alert-info">

The first step in creating an autograded assignment is to include the "Assignment Config" above.

The slides associated with the presentation provide a link to more details and options.

In order to add a question, we need to include the following:
- MarkDown Cell: Context and Instructions
- Raw Cell: Otter-Grader Question Fig
- Code Cell: Code to be completes and the SOLUTION marked with a comment
- Raw and Code Cells: Defining the tests to be run


</div>

```python
nba_player_data_2024 = Table.read_table("player_data_nba_2024.csv")
wnba_player_data_2024 = Table.read_table("https://github.com/ucb-ds/demo-datasets/raw/main/wnba_data.csv")
```

<div class="alert alert-block alert-success">

**Question 1.1 (2 points)** 

We are loading the WNBA and NBA player datasets in the previous cell.

- Compute and display the difference in the number of rows between the datasets

</div>


```python
player_diff = ...
player_diff
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q1_1 (points: [0, 0])
```
>>> 400 < player_diff < 500
True
```
```
>>> player_diff > 0
True
```
